Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

---
<div style="background-color: #c1f2a5">


# PS10

In this problem set, we are going to implement k-means clustering.
    
# Instructions



Remember to do your problem set in Python 3. Fill in `#YOUR CODE HERE`.

Unless we specify otherwise, make sure: 
- that all plots are scaled in such a way that you can see what is going on (while still respecting specific plotting instructions) 
- that the general patterns are fairly represented.
- to label all x- and y-axes, and to include a title.
    
**Test cases are here to help you debug your code, but passing them successfully is not a guarantee that your code is correct.**
    
</div>

```python
%matplotlib inline
import numpy as np
import matplotlib.pyplot as plt

```

An example observation from the categories cat, dog, and mop are presented below:
![](images/test.png)
We assume each observation can be represented as a pair of ($x,y$) coordinates, i.e., each object is represented in two-dimensional space. Suppose we have observed some observations from each type of object, but have lost the information as to which instance belongs to which type!

To try and recover this information we will use an unsupervised learning algorithm called _k-means clustering_. As you may recall from lecture, the $k$ here refers to how many types of clusters we think exist in the data, and the goal of the algorithm is to assign labels to the data points using their distance to the centers (or means) of the clusters. In other words, the data points should be assigned the category label of the cluster whose center is closest to the data point. For this particular problem, we assume $k=3$. After randomly initializing cluster centers,
the algorithm can be broken down into two alternating steps: 

1. Update the label assignments of the data points based on the nearest cluster centers
2. Update the positions of the cluster centers to reflect the updated assignments of data points.

Before you begin, load the data we will be using. For answering the questions in this problem set, use the `centers` loaded from the `X.npz`  file below (i.e., do NOT randomly initialize the values yourself - the autograder for this problem relies on a "stock" initialization).

<div class="alert alert-warning">
**N.B.** we use non-random initializations for the cluster centers to make tests letting you know if your code is correct feasible until Q4. Normally, cluster centers would be randomly initialized, and we will implement this in Q5. 
</div>

```python
data = np.load('data/X.npz')
X = data['X']
centers = data['centers'] 

print ('[x,y] coordinates of observations: \n' + str(X))
print ('\n[x,y] coordinates of cluster centers: \n' + str(centers))
```

---

## Q1.1 Distance [SOLO, 2pts]

First, we will need a function that gives us the distance between two points. We can use _Euclidean distance_ to compute the distance between two points ($x_1,y_1$) and ($x_2,y_2$). Recall **from the MDS examples** that Euclidean distance in $\mathbb{R}^2$ is calculated as:

$$
distance((x_1,y_1),(x_2,y_2)) = \sqrt{(x_1 - x_2)^{2} + (y_1 - y_2)^{2}}
$$

<div class="alert alert-success">
Complete the `distance` function below to calculate the euclidean distance between two points in $\mathbb{R}^2$. Copy-paste your code in gradescope.
</div>

<!-- grade_id: distance, points: 0, grade: false, solution: true -->
```python
def distance(a, b):
    """
    Returns the Euclidean distance between two points, 
    a and b, in R^2.
    
    Parameters
    ----------
    a, b : numpy arrays of shape (2,)
        The (x,y) coordinates for two points, a and b, 
        in R^2. E.g., a[0] is the x coordinate, 
        and a[1] is the y coordinate.
            
    Returns
    -------
    distance : float
        The Euclidean distance between a and b
    """
    # YOUR CODE HERE
```

```python
# add your own test cases here!

```

<!-- grade_id: test_distance, points: 0.5, grade: true, solution: false -->
```python
"""Check distances computes the correct values"""
from numpy.testing import assert_allclose

assert_allclose(distance(np.array([0.0, 0.0]), np.array([0.0, 1.0])), 1.0)
assert_allclose(distance(np.array([3.0, 3.0]), np.array([4.3, 5.0])), 2.3853720883753127)
assert_allclose(distance(np.array([130.0, -25.0]), np.array([0.4, 15.0])), 135.63244449614552)

print("Success!")
```

## Q1.2 Assignments [HELP, 3pts]
<div class="alert alert-success">Now, we will write a function to update the cluster that each point is assigned to. To do so, we will compute the distance of a point to the center of each cluster, and figure out which center is the closest one. Complete the `update_assignments` function to do this using your `distances` function. Copy-paste your code in gradescope.

The next cell provides a test case. If it does not return "Success!", your code is incorrect.</div>

<!-- grade_id: update_assignments, points: 0, grade: false, solution: true -->
```python
def update_assignments(num_clusters, X, centers):
    """
    Returns the cluster assignment (number) for each data point 
    in X, computed as the closest cluster center. Therefore, the 
    output is going to be an array with the length corresponding 
    to the number of datapoints.
    
    Parameters
    ----------
    num_clusters : int
        The number of disjoint clusters (i.e., k) in 
        the X
    
    X : numpy array of shape (m, 2)
        An array of m data points in R^2.
    
    centers : numpy array of shape (num_clusters, 2)
        The coordinates for the centers of each cluster
        
    Returns
    -------
    cluster_assignments : numpy array of shape (m,)
        An array containing the cluster label assignments 
        for each data point in X. Each cluster label is an integer
        between 0 and (num_clusters - 1). 
    """
    # YOUR CODE HERE
```

```python
# add your own test cases here!

```

<!-- grade_id: test_update_assignments, points: 1.5, grade: true, solution: false -->
```python
"""Check update_assignments computes the correct values"""
from nose.tools import assert_equal
from numpy.testing import assert_array_equal

# load the data
data = np.load('data/X.npz')
X = data['X']

# validate update_assignments using different values
actual = update_assignments(2, X, np.array([[3, 2], [1, 4]]))
expected = np.array([
    0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 
    0, 0, 0, 0, 0, 0, 0, 0, 1, 1, 1, 1, 0, 0, 1, 0, 0])

# is the output of the correct shape?
assert_equal(actual.shape[0], X.shape[0])

# are the cluster labels correct?
assert_array_equal(expected, actual)

# validate update_assignments using different values
actual = update_assignments(3, X[:int(X.shape[0]/2)], np.array([X[0], X[1], X[2]]))
expected = np.array([0, 1, 2, 2, 0, 2, 1, 2, 2, 2, 0, 0, 0, 0, 0])

# is the output of the correct shape?
assert_equal(actual.shape[0], X.shape[0] / 2)

# are the cluster labels correct?
assert_array_equal(expected, actual)

# check that it uses distance
old_distance = distance
del distance
try:
    update_assignments(2, X, np.array([[3, 2], [1, 4]]))
except NameError:
    pass
else:
    raise AssertionError("update_assignments does not call distance")
finally:
    distance = old_distance
    del old_distance

print("Success!")
```

---
## Q2. Update [HELP, 3pts] 

<div class="alert alert-success">Now, we need to do the next step of the clustering algorithm: recompute the cluster centers based on which points are assigned to that cluster. Recall that the new centers are simply the two-dimensional means of each group of data points in the cluster. A two-dimensional mean is calculated by simply finding the mean of the x coordinates and the mean of the y coordinates. Complete the `update_parameters` function to do this (check out the doc string to make sure you return the right output). Copy your code in gradescope. </div>

<!-- grade_id: update_parameters, points: 0, grade: false, solution: true -->
```python
def update_parameters(num_clusters, X, cluster_assignment):
    """
    Recalculates cluster centers running update_assignments.
    
    Parameters
    ----------
    num_clusters : int
        The number of disjoint clusters (i.e., k) in 
        the X
    
    X : numpy array of shape (m, 2)
        An array of m data points in R^2
    
    cluster_assignment : numpy array of shape (m,)
        The array of cluster labels assigned to each data 
        point as returned from update_assignments
    
    Returns
    -------
    updated_centers : numpy array of shape (num_clusters, 2)
        An array containing the new positions for each of 
        the cluster centers
    """
    # YOUR CODE HERE
```

```python
# add your own test cases here!

```

<!-- grade_id: test_update_parameters, points: 1.5, grade: true, solution: false -->
```python
"""Check update_parameters computes the correct values"""
from nose.tools import assert_equal
from numpy.testing import assert_allclose

# load the data
data = np.load('data/X.npz')
X = data['X']

# validate update_assignments using different values
cluster_assignment1 = np.array([
    0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 
    0, 0, 0, 0, 0, 0, 0, 0, 1, 1, 1, 1, 0, 0, 1, 0, 0])
actual = update_parameters(2, X, cluster_assignment1)
expected = np.array([[ 3.24286584,  2.71362623], [ 2.80577245,  4.07633606]])
assert_allclose(expected, actual)

cluster_assignment2 = np.array([0, 1, 2, 2, 0, 2, 1, 2, 2, 2, 0, 0, 0, 0, 0])
actual = update_parameters(3, X[:int(X.shape[0]/2)], cluster_assignment2)
expected = np.array([[ 3.4914304 ,  2.79181724], [ 3.03095255,  2.02958778], [ 2.86686881,  1.76070598]])
assert_allclose(expected, actual, rtol=1e-6)
    
print("Success!")
```

---
## Q3. Running k-means [3pts, SOLO]

At this stage you are ready to run the $k$-means clustering algorithm! The `k_means` function below will call your functions from Q1.2 and Q2 to run the k-means algorithm on the data points in `X`. Note that for this problem we assume that $k = 3$.

Take a look at the function `k_means` we're providing below, and run the next cell to call it.

```python
def k_means(num_clusters, X, centers, update_assignments, 
            update_parameters, n_iter=4, plotting = True):
    """
    Runs the k-means algorithm for n_iter iterations and plots
    the results.
    
    Parameters
    ----------
    num_clusters : int
        The number of disjoint clusters (i.e., k) to search for
        in X
    
    X : numpy array of shape (m, 2)
        An array of m data points in R^2
        
    centers : numpy array of shape (num_clusters, 2)
        The coordinates for the centers of each cluster
        
    update_assignments : function
        The function you completed in part A
    
    update_parameters : function
        The function you completed in part B
    
    n_iter : int (optional)
        The number of iterations to run the k_means update procedure.
        If not specified, defaults to 4.
    
    plotting : boolean (optional)
        plot the steps if true
        
    Returns
    -------
    cluster_assignments : numpy array of shape (m,) 
        The final label assignments for each of the data points in X
        
    centers : numpy array of shape (num_clusters, 2)
        The final cluster centroids in R^2 after running k-means
        
    fig: the figure object
    """
    if plotting:
        fig, ax = plt.subplots(2, n_iter,figsize=(30,15))
        fig.tight_layout(h_pad=2.0)
        ax = ax.flatten()

    for i in range(n_iter):
        # Step 1: Update cluster assignments
        cluster_assignments = \
            update_assignments(num_clusters, X, centers)
           
        if plotting:
            # plot data with colors corresponding to cluster assignments
            for j in range(X.shape[0]):
                if cluster_assignments[j] == 0:
                    ax[2*i].plot(X[j,0], X[j,1], 'r.')
                elif cluster_assignments[j] == 1:
                    ax[2*i].plot(X[j,0], X[j,1], 'b.')
                else:
                    ax[2*i].plot(X[j,0], X[j,1], 'g.')

            # plot the centers as stars with the associated color
            ax[2*i].plot(centers[0,0], centers[0,1], 'r*', markersize=10)
            ax[2*i].plot(centers[1,0], centers[1,1], 'b*', markersize=10)
            if len(centers)>2:
                ax[2*i].plot(centers[2,0], centers[2,1], 'g*', markersize=10)
            ax[2*i].set_title('Step 1: \nIteration ' + str(i+1))

            ax[2*i].set_xlim([2, 4.5])
            ax[2*i].set_ylim([1, 4.5])
            ax[2*i].set_xticks([]) 
            ax[2*i].set_yticks([])
        
        # Step 2: Update the cluster centers
        centers = \
            update_parameters(num_clusters, X, cluster_assignments)
        
        if plotting:
            # Plot data assignments with the updated center positions
            for j in range(X.shape[0]):
                if cluster_assignments[j] == 0:
                    ax[2*i+1].plot(X[j,0], X[j,1], 'r.')
                elif cluster_assignments[j] == 1:
                    ax[2*i+1].plot(X[j,0], X[j,1], 'b.')
                else:
                    ax[2*i+1].plot(X[j,0], X[j,1], 'g.')

            # Plot cluster centers as stars
            ax[2*i+1].plot(centers[0][0], 
                           centers[0][1], 
                           'r*', markersize=10)
            ax[2*i+1].plot(centers[1][0], 
                           centers[1][1], 
                           'b*', markersize=10)
            
            if len(centers)>2:
                ax[2*i+1].plot(centers[2][0], 
                               centers[2][1], 
                               'g*', markersize=10)
                
            ax[2*i+1].set_title('Step 2: \nIteration ' + str(i+1))

            ax[2*i+1].set_xlim([2, 4.5])
            ax[2*i+1].set_ylim([1, 4.5])
            ax[2*i+1].set_xticks([])
            ax[2*i+1].set_yticks([])
        
    if plotting:
        plt.show()
        return cluster_assignments, centers, fig
    else: 
        return cluster_assignments, centers
```

```python
# load the data
data = np.load('data/X.npz')
X = data['X']
centers = data['centers'] 

# run k-means
cluster_assignments, updated_centers, fig = k_means(3, X, centers, update_assignments, update_parameters, n_iter=4)

fig.savefig('PS10_Q3.png')
```

If the functions you completed above are working properly, you should see a figure containing a subplot of the output from steps (1) and (2) for four iterations of the algorithm. This plot should give you a sense of how the algorithm progresses over time. The data points are each assigned to one of three colors corresponding to their current cluster label. The cluster centers are plotted as stars.

<div class="alert alert-success"> Upload the screenshot of PS10_Q3.png in Gradescope. Describe what happened in the 4 2-step iterations in 2-3 sentences. Do you think we need more iterations in this particular case? Explain why in one sentence, in gradescope. </div>

YOUR ANSWER HERE.



## Q4. New object

Now that we have assigned cluster labels to each datapoint, let's investigate how we should classify a _new_ object (which we can see is a Shih-Tzu):

![](images/maddie.png)



Complete the function template in `assign_new_object` to determine the appropriate cluster label for this new object.

<!-- grade_id: assign_new_object, points: 0, grade: false, solution: true -->
```python
def assign_new_object(new_object, updated_centers):
    """
    Returns the cluster label (number) for new_object using k-means 
    clustering.
    
    Parameters
    ----------
    new_object : numpy array of shape (2,)
        The (x,y) coordinates of a new object to be classified
        
    updated_centers : numpy array of shape (num_clusters,2)
        An array containing the updated (x,y) coordinates for 
        each cluster center
        
    Returns
    -------
    label : int
       The cluster label assignment for new_object. This is a
       number between 0 and and (num_clusters - 1).
    """
    # YOUR CODE HERE

```

```python
# add your own test cases here!

```

<!-- grade_id: test_assign_new_object, points: 1, grade: true, solution: false -->
```python
"""Check assign_new_object computes the correct values"""
from nose.tools import assert_equal

# validate update_assignments using different values
centers1 = np.array([[ 3.17014624,  2.42738134], [ 2.90932354,  4.26426491]])
assert_equal(assign_new_object(np.array([0, 1]), centers1), 0)
assert_equal(assign_new_object(np.array([1, 0]), centers1), 0)
assert_equal(assign_new_object(np.array([3, 2]), centers1), 0)
assert_equal(assign_new_object(np.array([2, 4]), centers1), 1)

centers2 = np.array([[ 3.170146,  2.427381], [ 3.109456,  1.902395], [ 2.964183,  1.827484]])
assert_equal(assign_new_object(np.array([0, 1]), centers2), 2)
assert_equal(assign_new_object(np.array([1, 0]), centers2), 2)
assert_equal(assign_new_object(np.array([3, 2]), centers2), 1)
assert_equal(assign_new_object(np.array([2, 4]), centers2), 0)

# check that it uses distance
old_distance = distance
del distance
try:
    update_assignments(2, X, np.array([[3, 2], [1, 4]]))
except NameError:
    pass
else:
    raise AssertionError("assign_new_object does not call distance")
finally:
    distance = old_distance
    del old_distance

print("Success!")
```

Now that we have the function, let's first rerun $k$-means, to make sure we have the correct variables set:

```python
# load the edata
data = np.load('data/X.npz')
X = data['X']
centers = data['centers'] 

# run k-means
cluster_assignments, updated_centers, fig = k_means(3, X, centers, update_assignments, update_parameters, n_iter=4)
```

Next, let's implement the `assign_new_object` function to classify image of the Shih-Tzu:

```python
new_object = np.array([3.3, 3.5]) # image coordinates
label = assign_new_object(new_object, updated_centers)
print ('The new object was assigned to cluster: '+ str(label))
```

Finally, we can check the classification result against the true assignments using the provided helper function `plot_final`:

```python
def plot_final(X, cluster_assignments, updated_centers, new_object,
               assign_new_object):
    """
    Categorizes a new object and plots it against the true cluster
    labels.
    
    Parameters
    ----------    
    X : numpy array of shape (m, 2)
        An array of m data points in R^2
    
    cluster_assignments : numpy array of shape (m,) 
        The final label assignments for each of the data points in X
    
    updated_centers : numpy array of shape (num_clusters, 2)
        The coordinates for the centers of each cluster after 
        running k_means
        
    new_object : numpy array of shape (2,)
        The (x,y) coordinates of a new object to be classified
    
    assign_new_object : function
        The function you completed in part D    
        
    Returns
    -------
    label : fig
       The plotted figure.
    """
    fig, ax = plt.subplots(1,2)
    
    # plot data with colors corresponding to cluster assignments
    
    ax[0].plot(X[(cluster_assignments==0),0], X[(cluster_assignments==0),1], 'r.', label='0')
    ax[0].plot(X[(cluster_assignments==1),0], X[(cluster_assignments==1),1], 'b.', label='1')
    ax[0].plot(X[(cluster_assignments==2),0], X[(cluster_assignments==2),1], 'g.', label='2')
    
    # Generate a label for the new object
    label = assign_new_object(new_object, updated_centers)
    
    # Plot the new object as as big circle on the plot
    if label == 0:
        ax[0].plot(new_object[0], new_object[1], 'ro', markersize=10)
    elif label == 1:
        ax[0].plot(new_object[0], new_object[1], 'bo', markersize=10)
    else:
        ax[0].plot(new_object[0], new_object[1], 'go', markersize=10)
    
    ax[0].set_aspect('equal', 'datalim')
    ax[0].set_xticks([])
    ax[0].set_yticks([])
    ax[0].set_title('Final Cluster Assignments')
    ax[0].legend()
    
    # Plot the true cluster assignments for comparison
    ax[1].plot(X[:10, 0], X[:10, 1], 'r.', label='cats')
    ax[1].plot(X[10:20,0], X[10:20, 1], 'b.', label='dogs')
    ax[1].plot(X[20:, 0], X[20:, 1], 'g.', label='mops')
    ax[1].set_title('True Cluster Assignments')
    ax[1].set_aspect('equal', 'datalim')
    ax[1].set_xticks([])
    ax[1].set_yticks([])
    ax[1].legend()
    plt.show()
    
    return fig

```

```python
fig = plot_final(X, cluster_assignments, updated_centers, new_object, assign_new_object)

fig.savefig('PS10_Q4.png')
```


## Q4.1. New object [HELP, 2pts].

<div class="alert alert-success">
 When interpreting these plots, don't worry if the coloring differs between the two solutions; what matters is whether $k$-means identifies the same cluster boundaries as are shown in the true clusters. In gradescope, explain in one sentence why the coloring scheme does not matter - [hint: what do the cluster labels mean in the algorithm?]
    


<!-- grade_id: part_e1, points: 0.25, grade: true, solution: true -->
YOUR ANSWER HERE



## Q4.2. New object [SOLO, 5pts].

<div class="alert alert-success">
    
A. Did the algorithm correctly identify the Shih-Tzu?  Give an Yes/No answer, and a one sentence explanation for how you reached this conclusion.
    
B. Do you notice any differences between the true clusters and those identified via $k$-means? Write a few sentences commenting on any differences you found and why these differences might exist. Hint: are some dots that are originally in one cluster incorrectly identified as a part of a different cluster? Why might that be the case?
    
Upload your figure PS10_Q4.png in gradescope.</div>

<!-- grade_id: part_e1, points: 0.25, grade: true, solution: true -->
YOUR ANSWER HERE


## Q5. Random starting points [HELP, 5pts].

<div class="alert alert-success">
We are providing a function `init_centers` which will return random starting points for the center of each cluster. Take a look at the function below.  Re-run k-means as you did in Q3, but use init_centers to initialize the random starting points for the centers. Use 6 iterations, rather than 4, to run k-means clustering. Try  running it out multiple times, until your final cluster result is different from the one we obtained in Q3. Save and upload this figure in gradescope as PS10_Q5.png. Explain in 2-3 sentences how this solution differs from the original one. What happened in the algorithm to produce such results? </div>
    


```python

def init_centers(k=3):
    """
    Randomly initialize cluster centers
    """
    return np.vstack([np.random.uniform(low=2., high=4.5, size=k),
                      np.random.uniform(low=1, high=4.5, size=k)]).T
```

```python
# YOUR CODE HERE

fig. savefig('PS10_Q5.png')
```

YOUR ANSWER HERE.


## Q6. Numbers of clusters.

**So far, we have seen examples where we know that there are 3 categories/labels in the data. However, in general you won't know how many clusters there are in the data ahead of time. In this case, you need to figure out the correct/optimal number of clusters (k). Let's try to figure out a way to decide k.**

### Q6.1 k=2. [HELP, 2pts]

<div class="alert alert-success">
Run k-means with k=2, and 4 iterations. Use the cluster center starting points provided by the data (first two of the three). Plot final clusters and upload the figure as PS10_Q6_1.png to Gradescope. Explain the patterns you observe - how does clustering compare to previous results?
</div>
    

```python
# load the data
data = np.load('data/X.npz')
X = data['X']
centers = data['centers'][(0,2),:] 

# YOUR CODE HERE

#fig.savefig('PS10_Q6_1.png')
```

YOUR ANSWER HERE.

### Q6.2 Average distance. [SOLO, 3pts]

<div class="alert alert-success">
Write a function that computes the average distance a data point to the center of its cluster. We'll try use this measure to see how well the algorithm is doing. Enter in gradescope the average distance for the output of k-means with 4 iterations and the starting point centers provided above. Specifically, copy-paste the printed output of the cell (this will be auto-graded).</div>
    

```python

def average_distance(X,centers,cluster_assignments):
    """
    returns the average distance to clusters
    
    Parameters
    ----------    
    X : numpy array of shape (m, 2)
        An array of m data points in R^2
    
    cluster_assignments : numpy array of shape (m,) 
        The final label assignments for each of the data points in X
    
    centers : numpy array of shape (num_clusters, 2)
        The coordinates for the centers of each cluster after 
        running k_means 
        
    Returns
    -------
    label : 
       The average distance.
    """
    #YOUR CODE HERE
    
```

```python
# load the data
data = np.load('data/X.npz')
X = data['X']
centers = data['centers'] 

# run k-means
cluster_assignments, updated_centers, fig = k_means(3, X, centers, update_assignments, update_parameters, n_iter=4)

# compute the final distance
print("Final distance is: " + str(np.around(average_distance(X,updated_centers,cluster_assignments),decimals=2)))
```


### Q6.3 Number of clusters. [SOLO, 5pts]

<div class="alert alert-success">
Compute the average distance for k = 2 through 6, always using 10 iterations of the algorithm. For each value of k, compute the average distance for 20 different random starting point  initialized using init_centers function, and report the minimum over the starting points for this value of k. You should end up with a list or arrya of 5 minimum average distance values, one for each value of k. Create a bar plot, showing this minimum average distance as a function of k (x axis = k, y axis = minimum distance). Make sure to set up `plotting=False` as an argument in k-means, so you don't always plot the clustering results for each 20*6 times you run it! Upload your result as PS10_Q6_3.png in gradescope.</div>

```python
ks = np.arange(2,7)
nsims = 20
#YOUR CODE HERE


figure.savefig('PS10_Q6_3.png')

```

### Q6.4 Evaluating k-means. [HELP, 3]

If you got Q6.3 right, you should see that our average distance measure identifies $k=6$ as better than $k=3$. However, we know that's not the case: there are three clusters in our data. Therefore, our way of trying to determine the optimal number of clusters is flawed. In 2-3 sentences, explain why the measure we chose is not a good measure to determine which value of $k$ we should select. Hint: consider what would happen if we selected $k=$ the number of points in our data set. 

YOUR ANSWER HERE.

---
<div style="background-color: #c1f2a5">

# Submission

When you're done with your problem set, do the following:
- Upload your answers in Gradescope's PS5.

- Convert your Jupyter Notebook into a `.py` file by doing so:    
    
</div>


<center>    
  <img src="https://www.dropbox.com/s/7s189m4dsvu5j65/instruction.png?dl=1" width="300"/>
</center>

<div style="background-color: #c1f2a5">
    
- Submit the `.py` file you just created in Gradescope's PS5-code.
    
</div>        




</div>

</div>


```python

```